```python
#PLEASE DO NOT MODIFY THIS CELL

import unittest
from unittest.mock import patch
```

# First Semester Project
## Food Ordering Application
A food vendor has approached you build a simple app to help theirs users order food. 
The vendor sells the following food items
* Pizza at 6500 Naira
* Burger at 3000 Naira
* Noodles at 1300 Naira

Write the code for each of these functions to accurately capture a users order and display a receipt to the user.

NOTE: Users can order multiple items in different quantities

## INSTRUCTIONS
* Only type in the areas where you see the comment **"#Your Code Here"**.
* Delete this line **"raise NotImplementedError()"** you will find it directly under the **"#Your code here"** comment.
* Rename the file to the email you used in registering for this course.
* Do not delete any cells or change what is written in them except for the cells that you are supposed to write on.
* This project is worth 30 Marks so take it seriously and adhere strictly to the instructions.

```python
#PLEASE DO NOT MODIFY THIS CELL

def display_menu():
    """
    Description: Prints the menu options for the food items available in the ordering app.
                 since the app only has 3 food items you can order we add a fourth option to exit
                 the menu when they are done ordering.
    """
    print("Menu:")
    print("1. Pizza - 6500")
    print("2. Burger - 3000")
    print("3. Noodles - 1300")
    print("4. Exit Menu")
```

<!-- grade_id: cell-6c24609d20d6ba1a, points: 0, grade: false, solution: true -->
```python
def get_user_choice():
    """
    Description: Takes user input to get the number corresponding to the chosen food item from the menu.
                 Ensures the input is a valid choice between 1 and 4.

                 If the user types a valid number that is not between 1 and 4, return the ERROR
    """
    
    # using python exception to control the test output. It works with try and except in the code
    
    while True:
        try:
            choice = int(input("What are you buying today?, Select between 1 and 3 or 4 to Exit: "))
            
            if 1 <= choice <= 3: # returns as zero is less than 1 & it's not a valid option
                return choice
            
            elif choice == 4: # option 4 is to break the loop and quit
                # print("Thank you for choosing us! :)")
                return None
            
            else: # returns as choice is greater than given option & it's not a valid option
                print("Please enter a valid number!")
            
        except ValueError: # final python exception execution. if it encounters such, it will ignore and re-prompt
            print("Please enter a valid number")
```

<!-- grade_id: cell-3ddef8229d1713a5, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestGetUserChoice(unittest.TestCase):
    @patch('builtins.input', side_effect=['2'])
    def test_valid_choice(self, mock_input):
        result = get_user_choice()
        self.assertEqual(result, 2)

    @patch('builtins.input', side_effect=['invalid', '3'])
    def test_invalid_then_valid_choice(self, mock_input):
        result = get_user_choice()
        self.assertEqual(result, 3)

    @patch('builtins.input', side_effect=['5', '4'])
    def test_invalid_then_exit_choice(self, mock_input):
        result = get_user_choice()
        self.assertIsNone(result)

tester = TestGetUserChoice()
tester.test_valid_choice()
tester.test_invalid_then_valid_choice()
tester.test_invalid_then_exit_choice()
```

Output:
```
Please enter a valid number
Please enter a valid number!
```

<!-- grade_id: cell-c6a618a607e23f46, points: 0, grade: false, solution: true -->
```python
def get_quantity():
    """
    Description: Takes user input to get the quantity of the selected food item.
                 and ensures the input is a positive integer.
    """
    
    # using python exception to control the test output. It works with try and except in the code
    
    while True:
        try:
            quantity = int(input("How many do you wish to order? ")) # prompts quantity of choice
            if quantity > 0:
                return quantity
            else:
                print("Please enter a valid number!")
        except ValueError:
            print("Please enter a valid number!")
```

<!-- grade_id: cell-7fd228c806e327c0, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestGetQuantity(unittest.TestCase):
    @patch('builtins.input', side_effect=['3'])
    def test_valid_quantity(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 3)

    @patch('builtins.input', side_effect=['invalid', '5'])
    def test_invalid_then_valid_quantity(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 5)

    @patch('builtins.input', side_effect=['0', '-2', '4'])
    def test_invalid_then_valid_quantity_with_negative_input(self, mock_input):
        result = get_quantity()
        self.assertEqual(result, 4)

tester = TestGetQuantity()
tester.test_valid_quantity()
tester.test_invalid_then_valid_quantity()
tester.test_invalid_then_valid_quantity_with_negative_input()
```

Output:
```
Please enter a valid number!
Please enter a valid number!
Please enter a valid number!
```

<!-- grade_id: cell-fe851e407b6fa2ce, points: 0, grade: false, solution: true -->
```python
def get_item_name(choice):
    """
    Description: Retrieves and returns the name of a food item based on the user's choice from the menu.
    """
    
    
    names = {1: 'Pizza', 2: 'Burger', 3: 'Noodles'}
    return names.get(choice, "Unknown Item")
    # using get method to retrieve the name based on choice wit an "Unknown Item" prompt if not found
```

<!-- grade_id: cell-c061919667620632, points: 3, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

assert get_item_name(1) == 'Pizza'
assert get_item_name(2) == 'Burger'
assert get_item_name(3) == 'Noodles'
```

<!-- grade_id: cell-8266a8ef77b7a512, points: 0, grade: false, solution: true -->
```python
def get_item_price(choice):
    """
    Description: Retrieves and returns the price of a food item based on the user's choice from the menu.
    """
    
    
    prices = {1: 6500, 2: 3000, 3: 1300}
    return prices.get(choice, 0)
```

<!-- grade_id: cell-6b8ba807174ce818, points: 3, grade: true, solution: false -->
```python
assert get_item_price(1) == 6500
assert get_item_price(2) == 3000
assert get_item_price(3) == 1300
```

<!-- grade_id: cell-b5776935997d481b, points: 0, grade: false, solution: true -->
```python
def calculate_total_price(item_price, quantity):
    """
    Description: Calculates and returns the total price of a specific food item based on its price and the quantity ordered.
    """
    
    
    return item_price * quantity
```

<!-- grade_id: cell-2b5e0500570258af, points: 2, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

assert calculate_total_price(5, 2) == 10
```

<!-- grade_id: cell-02a3728e75ef8d90, points: 0, grade: false, solution: true -->
```python
def place_order():
    """
    Description: Manages the process of adding items to a shopping cart. 
                 USE A DICTIONARY FOR THE CART.
                 Calls other functions to get user choices, quantities, and calculates total prices.

                 Your cart should look something like this assuming this user ordered 3 pizzas and 3 burgers.
                 {
                    'Pizza': {'quantity': 3, 'total_price': 19500},
                    'Burger': {'quantity': 3, 'total_price': 9000}
                 }
    """
    
    # this programs calls all previously defined function
    
    cart = {}
    while True:
        display_menu()
        choice = get_user_choice()
        if choice is None:  # Exit option which breaks the loop as defined
            break
        quantity = get_quantity()
        total_price = calculate_total_price(get_item_price(choice), quantity)
        item_name = get_item_name(choice)
        
        if item_name in cart:
            cart[item_name]["quantity"] += quantity
            cart[item_name]["total_price"] += total_price
            
        else:
            cart[item_name] = {"quantity": quantity, "total_price": total_price}
            
        print(f"{quantity} {item_name}(s) added to your cart. Total price: {total_price}")

    return cart
```

<!-- grade_id: cell-16918c8d89ee10fa, points: 6, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestPlaceOrder(unittest.TestCase):
    @patch('__main__.get_user_choice', side_effect=[1, 2, None])
    @patch('__main__.get_quantity', return_value=3)
    def test_place_order(self, mock_get_quantity, mock_get_user_choice):
        result = place_order()

        # Assertions based on the expected behavior of place_order
        expected_result = {'Pizza': {'quantity': 3, 'total_price': 19500},
                           'Burger': {'quantity': 3, 'total_price': 9000}}
        self.assertEqual(result, expected_result)

        # Check that get_user_choice was called three times
        self.assertEqual(mock_get_user_choice.call_count, 3)

        # Check that get_quantity was called twice (for the two items added)
        self.assertEqual(mock_get_quantity.call_count, 2)

tester = TestPlaceOrder()
tester.test_place_order()
```

Output:
```
Menu:
1. Pizza - 6500
2. Burger - 3000
3. Noodles - 1300
4. Exit Menu
3 Pizza(s) added to your cart. Total price: 19500
Menu:
1. Pizza - 6500
2. Burger - 3000
3. Noodles - 1300
4. Exit Menu
3 Burger(s) added to your cart. Total price: 9000
Menu:
1. Pizza - 6500
2. Burger - 3000
3. Noodles - 1300
4. Exit Menu
```

<!-- grade_id: cell-620884e13f967fa3, points: 0, grade: false, solution: true -->
```python
def check_out(cart):
    """
    Description: Finalizes the order by displaying the contents of the shopping cart, including quantities and total prices.
                 Prints the total order price like a receipt.

                 The reciept would look like this if the cart is empty
                 
                     Your cart is empty. No items to check out.


                 If the Cart is has items in it then the receipt should look like this

                     Checking out...
                     Your order details:
                     Item 1: Quantity - 2, Total Price - $20
                     Item 2: Quantity - 3, Total Price - $15
                     Total Order Price: $35
                     Thank you for ordering!
    """
    
    # This program gives $ output based on the defined test. Please note that when we code in naira or other currenccy, it returns error
    
    
    if not cart:
        print("Your cart is empty. No items to check out.")
    else:
        print("Checking out...")
        print("Your order details:")
        total_order_price = 0
        for index, (item, details) in enumerate(cart.items(), start=1):
            print(f"Item {index}: Quantity - {details['quantity']}, Total Price - ${details['total_price']}")
            total_order_price += details['total_price']
        print(f"Total Order Price: ${total_order_price}")
        print("Thank you for ordering!")
```

<!-- grade_id: cell-db0a762b93d269b1, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestCheckOut(unittest.TestCase):
    @patch('builtins.print')
    def test_check_out_empty_cart(self, mock_print):
        cart = {}
        check_out(cart)
        mock_print.assert_called_with("Your cart is empty. No items to check out.")

    @patch('builtins.print')
    def test_check_out_non_empty_cart(self, mock_print):
        cart = {
            'Item 1': {'quantity': 2, 'total_price': 20},
            'Item 2': {'quantity': 3, 'total_price': 15}
        }
        check_out(cart)

        # Verify that the expected output was printed
        expected_output = [
            "Checking out...",
            "Your order details:",
            "Item 1: Quantity - 2, Total Price - $20",
            "Item 2: Quantity - 3, Total Price - $15",
            "Total Order Price: $35",
            "Thank you for ordering!"
        ]
        calls = [unittest.mock.call(output) for output in expected_output]
        mock_print.assert_has_calls(calls, any_order=False)

tester = TestCheckOut()
tester.test_check_out_empty_cart()
tester.test_check_out_non_empty_cart()
```

```python
#PLEASE DO NOT MODIFY THIS CELL

def food_ordering_app():
    """
    Description: The main function that initiates the food ordering application.
                 Calls place_order() to build the shopping cart and then calls check_out() to complete the order.

                 NOTE THAT IF ANY OF THE OTHER FUNCTIONS ARE NOT CORRECTLY WRITTEN THIS WILL FAIL
                 PLEASE DO NOT MODIFY THIS CELL
    """
    print("Welcome to the Food Ordering App!")
    cart = place_order()
    check_out(cart)
```

<!-- grade_id: cell-bc7e2ab74f6ce4b3, points: 4, grade: true, solution: false -->
```python
#PLEASE DO NOT MODIFY THIS CELL

class TestFoodOrderingApp(unittest.TestCase):
    @patch('builtins.print')
    @patch('__main__.place_order', return_value={'Item 1': {'quantity': 2, 'total_price': 20}})
    @patch('__main__.check_out')
    def test_food_ordering_app(self, mock_check_out, mock_place_order, mock_print):
        food_ordering_app()

        # Verify that the expected calls were made
        mock_print.assert_called_with("Welcome to the Food Ordering App!")
        mock_place_order.assert_called_once()
        mock_check_out.assert_called_once()

tester = TestFoodOrderingApp()
tester.test_food_ordering_app()
```

```python

```